<br>
<br>
<div>
<img  style="float: left; padding-right: 100px; width: 350px" src="logai.png">
    </div>
    <br>
    <h2 align="center">LogAI</h2>
<hr>
<h4 align="center"><a href="https://t.co/IoL5C0LiPX">Zephania Reuben</a></h4>
<h4 align="center">September, 2025</h4>
<hr>
<h3 align="center">AfriAI Lab - ADVANCED AI SHORT COURSE TRAINING 2025
<br>
<hr>
DS1: PYTHON FUNDAMENTALS | EXERCISE
</h3>
<hr>
<br>


```python
# Initialize Otter
import otter
grader = otter.Notebook("python_exercise.ipynb")
```

```python
# Check the IPython version
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

```python
# Import the good stuff
from nose.tools import assert_equal
import numpy as np
import math
import otter
grader = otter.Notebook()
```

**Question 1:**  Write a function `ypoly` to output the value of `y` given `x` and `p`. Function definition is `y = 4 + 5x^p` Ensure that if the `x` is not positive the value of `y` should be `0`

```python
def ypoly(x, p):
    ...
```

```python
grader.check("q1")
```

**Question 2:**  Write a function `perfectly_divisible` to return a list of numbers between 0 and a number $N$ that are perfectly divisible by $q$ (with out leaving a remainder). <br>
**Hint**: $N$ should also be inclusive in the numbers being considered. If N is negative use N == 20

```python
def perfectly_divisible(N, q):
    ...
```

```python
grader.check("q2")
```

**Question 3:**  Write a function `flatten_lists` that takes in a list of lists and outputs a sorted list of elements of sublists of the input list (confusing right?) <br>
Example: given `flatten_lists([[2,13,44], [6,7]])` it should return `[2,6,7,13,44]`

```python
def flatten_lists(superlist):
    ...
```

```python
grader.check("q3")
```

**Question 4:**  Write a function `longest_word` to return the longest string in a line of text <br>
Example: for text `Hallo people of the world!` should return `people` <br>
If there is a tie, then it should return the `1st` "longest" word encountered

```python
def longest_word(sentence):
    ...
```

```python
grader.check("q4")
```

**Question 5:**  Write a function `dictionarize` that takes a list as input and returns a dictionary with the elements of the list as the keys and the frequency of the element as the values.

```python
def dictionarize(alist):
    ...
```

```python
grader.check("q5")
```

**Question 6:**  Write a function `greatest_common_divisor` that takes two inputs `a` and `b` and returns the greatest common divisor of the two numbers. E.g. input `(10, 15)` would return `5`

```python
def greatest_common_divisor(a, b):
    ...
```

```python
grader.check("q6")
```

**Question 7:**  Write a function `get_nearest_farthest` that takes in a point of interest (pt) and a list of points and returns the nearest point and the farthest point from the point of intest.

```python
def get_nearest_farthest(pt, ptlist):
    ...
```

```python
grader.check("q7")
```

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

```python
# Save your notebook first, then run this cell to export your submission.
grader.export()
```

Credit: [Data Science Africa 2022](https://www.datascienceafrica.org/dsa2022arusha/blog/summer-school)

Tests for q1 (points: 6)
```
>>> ypoly(3,3) == 139
True
```
```
>>> ypoly(-2,3) == 0
True
```


Tests for q2 (points: 5)
```
>>> assert_equal(perfectly_divisible(10,2), [0,2,4,6,8,10])
```


Tests for q3 (points: 4)
```
>>> assert_equal(flatten_lists([[2,13,44], [6,7]]), [2, 6, 7, 13, 44])
```


Tests for q4 (points: 5)
```
>>> assert_equal(longest_word("Hello there people"), 'people')
```


Tests for q5 (points: 5)
```
>>> assert_equal(dictionarize([1,2,2,3,4,5]), {1: 1, 2: 2, 3: 1, 4: 1, 5: 1})
```


Tests for q6 (points: 5)
```
>>> assert_equal(greatest_common_divisor(10, 15), 5)
```


Tests for q7 (points: 10)
```
>>> assert_equal(get_nearest_farthest((3, 8), [(9, 3), (8,5), (7,6)]), ((7, 6), (9, 3)))
```
